---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Clustering with Multi-Dimension Data

```yaml
type: VideoExercise
key: 8faaa30407
xp: 50
```

`@projector_key`
d91f4543e9bba8946b450fbc3beaa543
